---
title: "gradethis examples"
output: learnr::tutorial
runtime: shiny_prerendered
tutorial:
  id: "demo-types"
  version: 1.0
---

```{r setup, include=FALSE}

library(learnr)
library(gradethis)

total <- function(m) {
  print(paste("Total mark is ",m))
}

mark <- function(m) {
  print(paste("Get mark: ",m))
}
```


## Intro

What is 1+1?
```{r ex1, exercise=TRUE}

```

```{r ex1-solution}
2
```


```{r ex1-check}
total(5)
grade_this({
  ls(all.names=TRUE)
  if (.result == .solution) {
    mark(5)
    
    pass("You got 5 marks!")
  }
  else if (exists("y", envir=.envir_result)){
    y <- get("y", envir = .envir_result)
    if (y==100){
      mark(2)
      fail("You got 2 marks only!")
    }
  }
  
  mark(0)
  fail("Wrong Answer.")
})

#grade_this_code()
```



```{r mc, echo=FALSE}
question("Which of the following interpretations of the bar charts above is NOT valid?",
  answer("Among characters with `Neutral` alignment, males are the most common.", message = "Nope, not quite!"),
  answer("Across all genders, `Bad` is the most common alignment.", correct = TRUE),
  answer("There are more male characters than female characters in this dataset.",  message = "Try again!"), 
  allow_retry = TRUE
)
```

```{r details}
question_text(
  "Student Name:",
  answer(NULL, correct = TRUE),
  incorrect = "Ok",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```



